---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 1e15c44df4
xp: 100
```

First steps in R for an interactive programming environment to teach and better learn R

`@instructions`
Start with your first steps in the software environment DataCamp a interactive software environment of R
First read Step: 1) Read your dataset that you want to manipulate in

`@hint`
Please use the function # read.csv() (reads comma separated columns)

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
APC <- read.csv
```

`@sct`
```{r}
Yeah your good
```
